# Pset 0-2: functions and getting help

**This notebook is adapted from the [Python](https://www.kaggle.com/learn/python) course.  You can reference the tutorial at [this link](https://www.kaggle.com/colinmorris/functions-and-getting-help).**

---
In the code cells below, replace `raise NotImplementedError` with your code.

Functions are powerful. Try writing some yourself.

# 1.

Complete the body of the following function according to its docstring.

HINT: Python has a built-in function `round`.

<!-- grade_id: cell-23427fd2239114dc, points: 0, grade: false, solution: true -->
```python
def round_to_two_places(num):
    """Return the given number rounded to two decimal places. 
    
    >>> round_to_two_places(3.14159)
    3.14
    """
    # Replace this body with your own code.
    # ("pass" is a keyword that does literally nothing. We used it as a placeholder
    # because after we begin a code block, Python requires at least one line of code)
    pass

    # YOUR CODE HERE
    # raise NotImplementedError()
    return round(num, 2)
```

<!-- grade_id: cell-5cd5c667b560b73a, points: 2, grade: true, solution: false -->
```python
# Verify the round_to_two_places function

assert round_to_two_places(3.14159) == 3.14

```

# 2.
The help for `round` says that `ndigits` (the second argument) may be negative.
What do you think will happen when it is? Try the following examples in the following cell.

```
round(566, -1)
round(566, -2)
round(566, -3)
```

```python
# Put your test code here
print(round(566, -1))
print(round(566, -2))
print(round(566, -3))
```

As you've seen, ndigits=-1 rounds to the nearest 10, ndigits=-2 rounds to the nearest 100 and so on. Where might this be useful? Suppose we're dealing with large numbers:

The area of Finland is 338,424 km²
The area of Greenland is 2,166,086 km²

We probably don't care whether it's really 338,424, or 338,425, or 338,177. All those digits of accuracy are just distracting. We can chop them off by calling round() with ndigits=-3:

The area of Finland is 338,000 km²
The area of Greenland is 2,166,000 km²

<!-- grade_id: cell-c83ebdfbd6be28a8, points: 0, grade: false, solution: true -->
```python
# To get some credit, complete the function below as per its docstring

def fn_round(num):
    '''Rounds num to the nearest 1000
    
    >>> fn_round(35858)
    '''
    pass

    # YOUR CODE HERE
    # raise NotImplementedError()
    return round(num, -3)
```

<!-- grade_id: cell-eb98d3f0832a8cd8, points: 2, grade: true, solution: false -->
```python
# Check the correctness of fn_round
assert fn_round(35858) == 36000

```

# 3.

In pset0-1, the candy-sharing friends Alice, Bob and Carol tried to split candies evenly. For the sake of their friendship, any candies left over would be smashed. For example, if they collectively bring home 91 candies, they'll take 30 each and smash 1.

Write in the cell below a simple function called `to_smash` that will calculate the number of candies to smash for *any* number of total candies. The funcation should be able to take two arguments. The first argument is required and is the total number of candies. The function can optionally take a second argument representing the number of friends the candies are being split between. **If no second argument is provided, it should assume 3 friends, as before.**

Be sure to write a docstring that reflects this behaviour.

<!-- grade_id: cell-3c54df3564e87a64, points: 0, grade: false, solution: true -->
```python
# Write a function as per the description above

# YOUR CODE HERE
# raise NotImplementedError()
def to_smash(num, pep=3):
    """Returns the remainder of num/pep
    This is to determine how many candies are to be smashed

    If pep is not passed in, pep defaults to 3
    """

    smash = num % pep
    return smash

```

<!-- grade_id: cell-02392ebab27b4956, points: 2, grade: true, solution: false -->
```python
# Check the correct functioning of to_smash

assert to_smash(91, 4) == 3
assert to_smash(91) == 1

```

# 4. (Optional)

It may not be fun, but reading and understanding error messages will be an important part of your Python career.

Each code cell below contains some commented buggy code. For each cell...

1. Read the code and predict what you think will happen when it's run.
2. Then uncomment the code and run it to see what happens. (**Tip**: You can highlight several lines and press `ctrl`+`/` to toggle commenting.)
3. Fix the code (so that it accomplishes its intended purpose without throwing an exception)

```python
round_to_two_places(9.9999)
```

```python
x = -10
y = 5
# Which of the two variables above has the smallest absolute value?
smallest_abs = min(abs(x), abs(y))
```

```python
def f(x):
    y = abs(x)
    return y

print(f(5))
```

# Submission

- Ensure that all the lines with `raise NotImplementedError` have been removed from the code cells above.
- To validate your work, run all the cells (`Cell->Run All`) and ensure there is no error.
- Then download your `pset0-2.ipynb` and submit it through SOMAS under the Assinment labeled `pset0-2`, which is under Wk3.