# Pset 0-4: Lists

This notebook has been adapted from [Python](https://www.kaggle.com/learn/python) course.  You can reference the tutorial at [this link](https://www.kaggle.com/colinmorris/lists).

---


Things get more interesting with lists. You'll apply your new knowledge to solve the questions below. 

# 1.

Complete the function below according to its docstring.

<!-- grade_id: cell-348fab73f0f8a851, points: 0, grade: false, solution: true -->
```python
def select_second(L):
    """Return the second element of the given list. If the list has no second
    element, return None.
    """
    
    # YOUR CODE HERE
    # raise NotImplementedError()

    return L[1] if len(L) >= 2 else None


```

<!-- grade_id: cell-a4af9b4ce9f446ae, points: 2, grade: true, solution: false -->
```python
# Check the select_second() function.
# Do not modify this cell. It is only for testing purposes.

trial_list = [1, 2, 3]
assert select_second(trial_list) == 2
trial_list = [1]
assert select_second(trial_list) == None

```

# 2.

You are analyzing sports teams.  Members of each team are stored in a list. The Coach is the first name in the list, the captain is the second name in the list, and other players are listed after that. 
These lists are stored in another list, which starts with the best team and proceeds through the list to the worst team last.  Complete the function below to select the **captain** of the worst team.

<!-- grade_id: cell-0f82506a8287c530, points: 0, grade: false, solution: true -->
```python
def losing_team_captain(teams):
    """Given a list of teams, where each team is a list of names, return the 2nd player (captain)
    from the last listed team
    """
    
    # YOUR CODE HERE
    # raise NotImplementedError()
    
    return teams[-1][1]

```

<!-- grade_id: cell-9503127318cd7710, points: 2, grade: true, solution: false -->
```python
# Chec kthe losing_team_captain() function
# This cell is only for testing. Do not modify its contents.

teams = [
    ['Coach1', 'captain1', 'player11', 'player12'],
    ['Coach2', 'captain2', 'player21', 'player22'],
    ['Coach3', 'captain3', 'player31', 'player32']
]
assert losing_team_captain(teams) == 'captain3'

```

# 3.

The next iteration of Mario Kart will feature an extra-infuriating new item, the *Purple Shell*. When used, it warps the last place racer into first place and the first place racer into last place. Complete the function below to implement the Purple Shell's effect.

<!-- grade_id: cell-9dc5959c375030b3, points: 0, grade: false, solution: true -->
```python
def purple_shell(racers):
    """Given a list of racers, set the first place racer (at the front of the list) to last
    place and vice versa.
    
    >>> r = ["Mario", "Bowser", "Luigi"]
    >>> purple_shell(r)
    >>> r
    ["Luigi", "Bowser", "Mario"]
    """
    
    # YOUR CODE HERE
    # raise NotImplementedError()

    racers[0], racers[-1] = racers[-1], racers[0]
    return

```

<!-- grade_id: cell-2416a27aea96139b, points: 2, grade: true, solution: false -->
```python
# Check purple_shell() function
# As usual, do not modify this cell

r = ["Mario", "Bowser", "Bros", "Luigi"]
purple_shell(r)
assert r == ["Luigi", "Bowser", "Bros", "Mario"]

```

# 4.

What are the lengths of the following lists? Fill in the variable `lengths` with your predictions. (Try to make a prediction for each list *without* just calling `len()` on it.)

<!-- grade_id: cell-7701c9c944b9b2f3, points: 0, grade: false, solution: true -->
```python
a = [1, 2, 3]
b = [1, [2, 3]]
c = []
d = [1, 2, 3][1:]

# Put your predictions in the list below. Lengths should contain 4 numbers, the
# first being the length of a, the second being the length of b and so on.
lengths = [3, 2, 0, 2]

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-a1ebfb7db9bd86af, points: 2, grade: true, solution: false -->
```python
# Check the lengths
# assert lengths == [len(a), len(b), len(c), len(d)]

# Do not modify this cell

```

# 5. <span title="A bit spicy" style="color: darkgreen ">🌶️</span>

We're using lists to record people who attended our party and what order they arrived in. For example, the following list represents a party with 7 guests, in which Adela showed up first and Ford was the last to arrive:

    party_attendees = ['Adela', 'Fleda', 'Owen', 'May', 'Mona', 'Gilbert', 'Ford']

A guest is considered 'fashionably late' if they arrived after at least half of the party's guests. However, they must not be the very last guest (that's taking it too far). In the above example, Mona and Gilbert are the only guests who were fashionably late.

Complete the function below which takes a list of party attendees as well as a person, and tells us whether that person is fashionably late.

<!-- grade_id: cell-42a97d083f9eef4f, points: 0, grade: false, solution: true -->
```python
def fashionably_late(arrivals, name):
    """Given an ordered list of arrivals to the party and a name, return whether the guest with that
    name was fashionably late.
    """
    
    # YOUR CODE HERE
    # raise NotImplementedError()

    for i in arrivals[(len(arrivals)+1)//2 : -1]:
        if i == name:
            return True
        
    return False

```

<!-- grade_id: cell-3a8c7eb49cb52734, points: 4, grade: true, solution: false -->
```python
# Check the fashionably_late() function
# Kama kawaida, do not modify this cell

party_attendees = ['Adela', 'Fleda', 'Owen', 'May', 'Mona', 'Gilbert', 'Ford']
assert fashionably_late(party_attendees, 'Adela') == False
assert fashionably_late(party_attendees, 'May') == False
assert fashionably_late(party_attendees, 'Mona') == True
assert fashionably_late(party_attendees, 'Ford') == False

```

# Submission

- Be sure to remove all the `raise NotImplementedError` lines in all the cells above.
- Run all the cells (`Cells->Run All`) and ensure that there is no error.
- Download your `pset0-4.ipynb` and submit through SOMAS under the Assignment labeled pset0-4, under Wk 4.