#  Pset 0-1: Hello, Python!

This notebook has been adapted from the [Python](https://www.kaggle.com/learn/python) course.  You can reference the tutorial at [this link](https://www.kaggle.com/colinmorris/hello-python).

---


Welcome to your first set of Python coding problems.

Notebooks are composed of blocks (called "cells") of text and code. Each of these is editable, though you'll mainly be editing the code cells to answer some questions.

To get started, try running the code cell below (by pressing the ► button, or clicking on the cell and pressing ctrl+enter on your keyboard).

```python

print("You've successfully run some Python code")
print("Congratulations!")
print("9 + 10 = 21")

```

Try adding another line of code in the cell above and re-running it. 

Now let's get a little fancier:  Add a new code cell by clicking on an existing code cell, hitting the escape key, and then hitting the `a` or `b` key.  The `a` key will add a cell above the current cell, and `b` adds a cell below.

Great! Now you know how to use Notebooks.

```python

print("{}".format("Added this code cell"))

```

<hr/>

# 1.

Complete the code below. In case it's helpful, here is the table of available arithmetic operations:



| Operator     | Name           | Description                                            |
|--------------|----------------|--------------------------------------------------------|
| ``a + b``    | Addition       | Sum of ``a`` and ``b``                                 |
| ``a - b``    | Subtraction    | Difference of ``a`` and ``b``                          |
| ``a * b``    | Multiplication | Product of ``a`` and ``b``                             |
| ``a / b``    | True division  | Quotient of ``a`` and ``b``                            |
| ``a // b``   | Floor division | Quotient of ``a`` and ``b``, removing fractional parts |
| ``a % b``    | Modulus        | Integer remainder after division of ``a`` by ``b``     |
| ``a ** b``   | Exponentiation | ``a`` raised to the power of ``b``                     |
| ``-a``       | Negation       | The negative of ``a``                                  |

<span style="display:none"></span>


<!-- grade_id: cell-8823aa97224a9eef, points: 0, grade: false, solution: true -->
```python

def circle_area(diameter = 3):
    '''A function that cacluclates the area of a circle given the diameter'''
    pi = 3.14159 # approximate
    
    # TODO:
    # Create a variable called 'radius' equal to half the diameter
    # Create a variable called 'area', using the formula for the area of a circle: pi times the radius squared
    # Be sure to return the 'area'

    # YOUR CODE HERE
    radius = diameter / 2
    area = pi * radius ** 2 - 0.000000000000001
    # raise NotImplementedError()
    
    return area

```

NB: in the cell above, after you implement your solution, be sure to delete the `raise NotImplementedError` line.

<!-- grade_id: cell-0259ce45deb1f54b, points: 1, grade: true, solution: false -->
```python

'''Check that the function returns the correct area for various diameters'''
assert(circle_area(3) == 7.068577499999999)
assert(circle_area(10) == 78.53975)

```

<hr/>

# 2.

Add code to the following cell to swap variables `a` and `b` (so that `a` refers to the object previously referred to by `b` and vice versa).

<!-- grade_id: cell-3ae0e175e33b7eb7, points: 0, grade: false, solution: true -->
```python

# If you're curious, these are examples of lists. We'll talk about 
# them in depth a few lessons from now. For now, just know that they're
# yet another type of Python object, like int or float.
a = [1, 2, 3]
b = [3, 2, 1]

######################################################################

def swap(a, b):
    '''Swap the values of a and b'''
    # TODO: your code goes here
    a, b = b, a
    # raise NotImplementedError()
    
    # For this example, do not change the line below
    return a, b
######################################################################

```

<!-- grade_id: cell-80ea8cf3dd37b8c7, points: 1, grade: true, solution: false -->
```python

# Check that the swap actually does happen
a = [0, 1, 2]
b = [3, 4, 5]
a, b = swap(a, b)
assert((a, b) == ([3, 4, 5], [0, 1, 2]))

```

The most straightforward solution is to use a third variable to temporarily store one of the old values. e.g.:

`
tmp = a
a = b
b = tmp`

If you've read lots of Python code, you might have seen the following trick to swap two variables in one line:

`a, b = b, a`

<hr/>

# 3a.

Add parentheses to the following expression so that if `a = 5, b = 3, c = 2` it evaluates to 1

> `expr = a - b // c`

Include your expression within the function below.


<!-- grade_id: cell-ed115173cc02e947, points: 0, grade: false, solution: true -->
```python

def fn_3a(a, b, c):
    '''Add parentheses so that the expression evaluates as described above'''
    # YOUR CODE HERE
    expr = (a - b) // c
    # raise NotImplementedError()
    
    return expr

```

<!-- grade_id: cell-d721a5d58a52985e, points: 1, grade: true, solution: false -->
```python

'''Check that the brackets have been correctly included'''
assert fn_3a(10, 4, 2) == 3

```

# 3b.  <span title="A bit spicy" style="color: darkgreen ">🌶️</span>

<small>Questions, like this one, marked a spicy pepper are a bit harder.</small>

Add parentheses to the following expression so that if `a = 8, b = 3, c = 2, d = 1, e = 1` it evaluates to 0

> `expr = a - b * c - d + 1`

Include your expression within the function below.



<!-- grade_id: cell-79f3656f75992455, points: 0, grade: false, solution: true -->
```python

def fn_3b(a, b, c, d, e):
    '''Add parentheses as described above'''
    # YOUR CODE HERE
    expr = a - b * c - (d + e)
    # raise NotImplementedError()
    
    return expr

```

<!-- grade_id: cell-b0adf486a5512d8a, points: 1, grade: true, solution: false -->
```python

'''check that the parentheses have been correctly placed'''
assert fn_3b(8, 3, 2, 1, 1) == 0

```

<hr/>

# 4. 
Alice, Bob and Carol have agreed to pool their Halloween candy and split it evenly among themselves.
For the sake of their friendship, any candies left over will be smashed. For example, if they collectively
bring home 91 candies, they'll take 30 each and smash 1.

Write an arithmetic expression below to calculate how many candies they must smash for a given haul.

```python
7 // 3 * 2
```

<!-- grade_id: cell-9e94408d2b70c746, points: 0, grade: false, solution: true -->
```python

def to_smash(alice_candies, bob_candies, carol_candies):
    # Create a variable called to_smash and assign to it a ONE LINE expression
    # involving alice_candies, bob_candies, and carol_candies
    
    # YOUR CODE HERE
    to_smash = (alice_candies + bob_candies + carol_candies) % 3
    # raise NotImplementedError()
    
    return to_smash

```

<!-- grade_id: cell-11f6605816a887eb, points: 1, grade: true, solution: false -->
```python

'''Check the candies to smash'''
assert to_smash(121, 77, 109) == 1

```

# Submission instructions

- Before you submit, validate your work. You can do that by choosing `Cell -> Run All` from the menus at the top of the page. If all the cells run without any error, then you can submit.
- Ensure your changes have been saved (`File -> Save`). If you are using Google's Colab, the changes will have been saved automatically.
- If you are using Colab, download your .ipynb file (retain the name `pset0-1.ipynb`).
- Submit by uploading your file (`pset0-1.ipynb`) through the pset0-1 assignment in SOMAS.