# Pset 0-5: Loops and List Comprehensions

This notebook is adapted from [Python](https://www.kaggle.com/learn/python) course.  You can reference the tutorial at [this link](https://www.kaggle.com/colinmorris/loops-and-list-comprehensions).

---


With all you've learned, you can start writing much more interesting programs. See if you can solve the problems below.

# 1.

Have you ever felt debugging involved a bit of luck? The following program has a bug. Try to identify the bug and fix it.

Try to identify the bug and fix it in the cell below:

<!-- grade_id: cell-1e77d78e26c3777c, points: 0, grade: false, solution: true -->
```python
def has_lucky_number(nums):
    """Return whether the given list of numbers is lucky. A lucky list contains
    at least one number divisible by 7.
    """
    for num in nums:
        if num % 7 == 0:
            return True

    return False

# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-308b27365ad2ef77, points: 3, grade: true, solution: false -->
```python
# Verifying the cell above
# Do not modify this cell

assert has_lucky_number([2, 3, 7, 14]) == True
assert has_lucky_number([2, 3, 17]) == False

```

# 2.
Look at the Python expression below. What do you think we'll get when we run it? When you've made your prediction, uncomment the code and run the cell to see if you were right.

```python
# [1, 2, 3, 4] > 2

```

R and Python have some libraries (like numpy and pandas),  which can compare each element of the list to 2 (i.e. do an 'element-wise' comparison) and give us a list of booleans like `[False, False, True, True]`. 

Implement a function that reproduces this behaviour, returning a list of booleans corresponding to whether the corresponding element is greater than n. **If you use list comprehensions effectively, you should be able to give a one-line solution.**

<!-- grade_id: cell-275423a51f62e293, points: 0, grade: false, solution: true -->
```python
def elementwise_greater_than(L, thresh):
    """Return a list with the same length as L, where the value at index i is 
    True if L[i] is greater than thresh, and False otherwise.
    
    >>> elementwise_greater_than([1, 2, 3, 4], 2)
    [False, False, True, True]
    """
    
    # YOUR CODE HERE
    # raise NotImplementedError()
    bools = []
    for i in L:
        bools.append(True) if i > thresh else bools.append(False)

    return bools
```

<!-- grade_id: cell-69fb39a402cb3301, points: 3, grade: true, solution: false -->
```python
# Check the above solution
# Do not modify this cell

assert elementwise_greater_than([1, 2, 3, 4], 2) == [False, False, True, True]
```

# 3.

Complete the body of the function below according to its docstring.

<!-- grade_id: cell-b53c7bf2a168b7af, points: 0, grade: false, solution: true -->
```python
def menu_is_boring(meals):
    """Given a list of meals served over some period of time, return True if the
    same meal has ever been served two days in a row, and False otherwise.
    """
    
    # YOUR CODE HERE
    # raise NotImplementedError()

    no_rep = set(meals)
    return True if len(meals)!=len(no_rep) else False

```

<!-- grade_id: cell-e0bf09958333fe41, points: 3, grade: true, solution: false -->
```python
# Check solution for above cell

menu_1 = ['ugali', 'chips', 'githeri', 'matumbo', 'fish', 'tortilla', 'lassagna', 'churros', 'paella']
assert menu_is_boring(menu_1) == False
menu_2 = ['ugali', 'chips', 'githeri', 'matumbo', 'matumbo', 'tortilla', 'lassagna', 'churros', 'paella']
assert menu_is_boring(menu_2) == True

```

# Submission

- Be sure to remove all the `raise NotImplementedError` lines in all the cells above.
- Run all the cells (`Cells->Run All`) and ensure that there is no error.
- Download your `pset0-5.ipynb` and submit through SOMAS under the Assignment labeled pset0-5, under Wk 4.