# Pset 0-3: Functions and Conditionals

**This notebook is adapted from [Python](https://www.kaggle.com/learn/python) course.  You can reference the tutorial at [this link](https://www.kaggle.com/colinmorris/booleans-and-conditionals).**

---


**In the code cells below, replace `raise NotImplementedError` with your code.**

In this exercise, you'll put to work what you have learned about booleans and conditionals.

# 1.

Many programming languages have [`sign`](https://en.wikipedia.org/wiki/Sign_function) available as a built-in function. Python doesn't, but we can define our own!

In the cell below, define a function called `sign` which takes a numerical argument and returns -1 if it's negative, 1 if it's positive, and 0 if it's 0.

<!-- grade_id: cell-cf2705bc462b2070, points: 0, grade: false, solution: true -->
```python
# Your code goes here. Define a function called 'sign'

# YOUR CODE HERE
# raise NotImplementedError()
def sign(num):
    """Check if a number num is positive, negative or zero

    returns -1 if num is negative, 1 if it's positive and 0 if it equals zero
    """
    if num < 0:
        return -1

    elif num > 0:
        return 1
    
    return 0
    

```

<!-- grade_id: cell-d1e915e790ac2614, points: 2, grade: true, solution: false -->
```python
# Check the sign function above
assert sign(5) == 1
assert sign(-5) == -1
assert sign(0) == 0

```

# 2.

We've decided to add "logging" to our `to_smash` function from the previous exercise.

```python
def to_smash(total_candies):
    """Return the number of leftover candies that must be smashed after distributing
    the given number of candies evenly between 3 friends.
    
    >>> to_smash(91)
    1
    """
    print("Splitting", total_candies, "candies")
    return total_candies % 3

to_smash(91)
```

What happens if we call it with `total_candies = 1`?

```python
to_smash(1)
```

That isn't great grammar!

Modify the definition in the cell below to correct the grammar of our print statement. (If there's only one candy, we should use the singular "candy" instead of the plural "candies")

<!-- grade_id: cell-c98ce72e6f46e816, points: 0, grade: false, solution: true -->
```python
def to_smash(total_candies):
    """Return the number of leftover candies that must be smashed after distributing
    the given number of candies evenly between 3 friends.
    
    >>> to_smash(91)
    1
    """
    
    # YOUR CODE HERE
    # raise NotImplementedError()
    # Uncomment the line below and then adjust it accordingly

    candy = "candy" if total_candies % 3 == 1 else "candies"

    print("Splitting", total_candies, candy)

    return total_candies % 3


```

<!-- grade_id: cell-2ac3c98a7ae582e2, points: 2, grade: true, solution: false -->
```python
# Check the function
to_smash(91)
to_smash(1)

```

# 3. <span title="A bit spicy" style="color: darkgreen ">🌶️</span>

In the tutorial linked to at the top of this page, there is mention about deciding whether we're prepared for the weather. It was indicated that I'm safe from today's weather if...
- I have an umbrella...
- or if the rain isn't too heavy and I have a hood...
- otherwise, I'm still fine unless it's raining *and* it's a workday

The function below is an attempt at turning this logic into a Python expression. There is a bug in the code. Can you find it?

```
def prepared_for_weather(have_umbrella, rain_level, have_hood, is_workday):
    return have_umbrella or rain_level < 5 and have_hood or not rain_level > 0 and is_workday
```

Define the corrected version of the function in the cell below.

<!-- grade_id: cell-1d77fb5228179b9d, points: 0, grade: false, solution: true -->
```python
# Corrected version of prepared_for_weathere

def prepared_for_weather(have_umbrella, rain_level, have_hood, is_workday):
    # YOUR CODE HERE
    # raise NotImplementedError()
    # Uncomment the line below and then correct the bug in it.
    if rain_level == 0 or have_umbrella or (rain_level < 5 and have_hood) or not is_workday:
        return True

    return False

```

<!-- grade_id: cell-aa46f56fc11a7710, points: 4, grade: true, solution: false -->
```python
# Check that the prepared_for_weather function has been corrected
# Do not modify this cell

have_umbrella = True
rain_level = 0.0
have_hood = True
is_workday = True
assert prepared_for_weather(have_umbrella, rain_level, have_hood, is_workday) == True

have_umbrella = False
rain_level = 0.0
have_hood = False
is_workday = False
assert prepared_for_weather(have_umbrella, rain_level, have_hood, is_workday) == True

```

# 4.

The function `is_negative` below is implemented correctly - it returns True if the given number is negative and False otherwise.

However, it's more verbose than it needs to be. We can actually reduce the number of lines of code in this function by *75%* while keeping the same behaviour. 

See if you can come up with an equivalent body that uses just **one line** of code, and put it in the function `concise_is_negative`. (HINT: you don't even need Python's ternary syntax)

<!-- grade_id: cell-bca3586c56316591, points: 0, grade: false, solution: true -->
```python
def is_negative(number):
    if number < 0:
        return True
    else:
        return False

def concise_is_negative(number):
    # Keep your code to one line!)
    
    # YOUR CODE HERE
    # raise NotImplementedError()
    return True if number < 0 else False

```

<!-- grade_id: cell-201f1eceb720c5c3, points: 2, grade: true, solution: false -->
```python
# Check concise_is_negative

assert concise_is_negative(3) == False
assert concise_is_negative(-3) == True

```

# 5a.

The boolean variables `ketchup`, `mustard` and `onion` represent whether a customer wants a particular topping on their hot dog. We want to implement a number of boolean functions that correspond to some yes-or-no questions about the customer's order. For example:

```python
def onionless(ketchup, mustard, onion):
    """Return whether the customer doesn't want onions.
    """
    return not onion
```

<!-- grade_id: cell-f9db35d93ba049ff, points: 0, grade: false, solution: true -->
```python
def wants_all_toppings(ketchup, mustard, onion):
    """Return whether the customer wants "the works" (all 3 toppings)
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    if ketchup == mustard == onion == True:
        return True

    return False

```

<!-- grade_id: cell-0c671c34925088c7, points: 2, grade: true, solution: false -->
```python
# Check the wants_all_toppings function

assert wants_all_toppings(True, True, True) == True
assert wants_all_toppings(True, True, False) == False
assert wants_all_toppings(True, False, True) == False

```

# 5b.

For the next function, fill in the body to match the English description in the docstring. 

<!-- grade_id: cell-2482e1dd07d8b143, points: 0, grade: false, solution: true -->
```python
def wants_plain_hotdog(ketchup, mustard, onion):
    """Return whether the customer wants a plain hot dog with no toppings.
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    if ketchup == mustard == onion == False:
        return True
    
    return False

```

<!-- grade_id: cell-bd13042fa3d7268b, points: 2, grade: true, solution: false -->
```python
# Check the wants_plain_hotdog function

assert wants_plain_hotdog(False, False, False) == True
assert wants_plain_hotdog(True, False, False) == False
assert wants_plain_hotdog(False, False, True) == False
```

# 5c.

You know what to do: for the next function, fill in the body to match the English description in the docstring.

<!-- grade_id: cell-32a63a9d220d54cd, points: 0, grade: false, solution: true -->
```python
def exactly_one_sauce(ketchup, mustard, onion):
    """Return whether the customer wants either ketchup or mustard, but not both.
    (You may be familiar with this operation under the name "exclusive or")
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    
    return True if ketchup != mustard else False

```

<!-- grade_id: cell-2e08e93918937951, points: 2, grade: true, solution: false -->
```python
# Check the exactly_one_sauce function

assert exactly_one_sauce(True, False, True) == True
assert exactly_one_sauce(False, True, True) == True
assert exactly_one_sauce(True, True, True) == False
assert exactly_one_sauce(False, False, True) == False

```

# 6. <span title="A bit spicy" style="color: darkgreen ">🌶️</span>

We’ve seen that calling `bool()` on an integer returns `False` if it’s equal to 0 and `True` otherwise. What happens if we call `int()` on a bool? Try it out in the notebook cell below.



```python
print(int(True))
print(int(False))
```

Can you take advantage of this to write a succinct function that corresponds to the English sentence "does the customer want exactly one topping?"?

<!-- grade_id: cell-1b1187d56cd73ce0, points: 0, grade: false, solution: true -->
```python
def exactly_one_topping(ketchup, mustard, onion):
    """Return whether the customer wants exactly one of the three available toppings
    on their hot dog.
    """
    # YOUR CODE HERE
    # raise NotImplementedError()
    if int(ketchup) + int(mustard) + int(onion) == 1:
        return True
    
    return False
```

<!-- grade_id: cell-bf03f955eb501b20, points: 0, grade: true, solution: false -->
```python
# Check the exactly_one_topping function

assert exactly_one_topping(True, False, False) == True
assert exactly_one_topping(False, True, False) == True
assert exactly_one_topping(False, False, True) == True
assert exactly_one_topping(True, False, True) == False

```

# Submission

- Be sure to remove all the `raise NotImplementedError` lines in all the cells above.
- Run all the cells (`Cells->Run All`) and ensure that there is no error.
- Download your `pset0-3.ipynb` and submit through SOMAS under the Assignment labeled pset0-3, under Wk 3.